```python
import pandas as pd
import matplotlib.pyplot as plt
import numpy as np
import scipy
import scipy.stats
```

```python
# dataset source https://finance.yahoo.com/quote/%5ESP500TR/history?period1=568278000&period2=1548226800&interval=1d&filter=history&frequency=1d
# historical daily prices for S&P 500 TR (total return) index
# NOTE: Adj Close factors in dividends, so that's a more accurate measure of returns
url="https://raw.githubusercontent.com/marilynwaldman/kurtosis/master/SP500TR.csv"
df = pd.read_csv(url)
df.head()
```

```python
# take Closing value (today) - Closing value (yesterday)
returns = np.ediff1d(df['Adj Close'].values)  
returns

```

```python
# convert to % by dividing by Closing value (yesterday)
pct_returns = returns / df['Adj Close'].values[:-1]  # convert to % by dividing by Closing value (yesterday)
pct_returns
```

#### use np.min(array name) to compute the minimum 
np.min(array_name)

pct_rtns_min = np.min(pct_returns)


<!-- grade_id: cell-732b042a572f851c, points: 0, grade: false, solution: true -->
```python
#pct_rtns_min = np.min(pct_returns)
# YOUR CODE HERE
raise NotImplementedError()

print(pct_rtns_min)
```

<!-- grade_id: cell-3137297aaa7e11cd, points: 5, grade: true, solution: false -->
```python
assert pct_rtns_min == np.min(pct_returns)

```

#### use np.max(array name) to compute the minimum 
np.max(array_name)

pct_rtns_max = np.max(pct_returns)

<!-- grade_id: cell-5185ab4cf5d46fa9, points: 0, grade: false, solution: true -->
```python
# pct_rtns_max = np.max(pct_returns)
# YOUR CODE HERE
raise NotImplementedError()

print(pct_rtns_max)
```

<!-- grade_id: cell-388d499748f0db48, points: 5, grade: true, solution: false -->
```python
assert pct_rtns_max == np.max(pct_returns)
```

```python
# plot histogram compared to Gaussian.  Looks like this is not a Gaussian!
# in fact it is probably a Laplace distribution
# https://sixfigureinvesting.com/2016/03/modeling-stock-market-returns-with-laplace-distribution-instead-of-normal/
plt.figure(1)
plt.hist(pct_returns, bins=100, density=True)
plt.xlim((min(pct_returns), max(pct_returns)))

# calculate fitted Gaussian
mean = np.mean(pct_returns)
variance = np.var(pct_returns)
sigma = np.sqrt(variance)
x = np.linspace(min(pct_returns), max(pct_returns), 100)
plt.plot(x, scipy.stats.norm.pdf(x, mean, sigma))

plt.show()
```

```python
#### Compute the kurtosis - kurtosis(array_name)
```

<!-- grade_id: cell-942644d996cbbe26, points: 0, grade: false, solution: true -->
```python
# returns have LONG tails.  "Outlier" evens definitely happen (i.e. catastrophic market crashes)
# kurtosis=3 would be Gaussian.  Anything larger has fat tails
from scipy.stats import kurtosis
#pct_rtns_kurtosis = kurtosis(pct_returns)
# YOUR CODE HERE
raise NotImplementedError()

print(pct_rtns_kurtosis)
```

<!-- grade_id: cell-20f4ddd17535e2f6, points: 5, grade: true, solution: false -->
```python
assert pct_rtns_kurtosis == kurtosis(pct_returns)
```

```python
# plot again on log scale to really show those fat tails

plt.figure(1)
plt.hist(pct_returns, bins=20, density=True)
plt.xlim((min(pct_returns), max(pct_returns)))

# calculate fitted Gaussian
mean = np.mean(pct_returns)
variance = np.var(pct_returns)
sigma = np.sqrt(variance)
x = np.linspace(min(pct_returns), max(pct_returns), 100)
plt.plot(x, scipy.stats.norm.pdf(x, mean, sigma))

plt.yscale('log')
plt.show()
```

```python

```